<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Week 11 - Q5: Self-Inductance and Inductor Energy             -->
<!-- Course: Electromagnetism and Circuit Analysis, LUT University  -->
<!-- Topics: Self-inductance, solenoid, energy stored in inductor  -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Week 11 – Q5: Self-Inductance of a Solenoid and Inductor Energy</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Week 11 &ndash; Q5: Self-Inductance of a Solenoid and Inductor Energy</h3>
<p><em>Course: Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><em>This is a training exercise &mdash; you may use books, tutors, and AI. The question requires multi-step reasoning.</em></p>
<p><strong>Total: 4 points | Penalty per wrong attempt: 0%</strong></p>

<hr />

<p>A solenoid has \(N = {@N_turns@}\) turns, radius \(r = {@r_cm@}\;\text{cm}\),
length \(\ell = {@ell_cm@}\;\text{cm}\), and carries a current \(I = {@I_val@}\;\text{A}\).
Use \(\mu_0 = 4\pi \times 10^{-7}\;\text{H/m}\).</p>

<p style="text-align:center;">[DIAGRAM: Solenoid with N turns, radius r, length &ell;, current I]</p>

<hr />

<h4>Part (a) &mdash; Self-Inductance (1 point)</h4>
<p>Compute the self-inductance \(L\) of the solenoid using \(L = \mu_0 N^2 A / \ell\).
Give your answer in <strong>millihenries (mH)</strong>.</p>
<p>\(L = \) [[input:ans1]] mH [[validation:ans1]]</p>
<p><em>Syntax hint: enter a decimal number, e.g. <code>0.395</code>.</em></p>

<h4>Part (b) &mdash; Stored Energy (1 point)</h4>
<p>Compute the energy stored in the inductor using \(W = \tfrac{1}{2} L I^2\).
Give your answer in <strong>millijoules (mJ)</strong>.</p>
<p>\(W = \) [[input:ans2]] mJ [[validation:ans2]]</p>
<p><em>Syntax hint: enter a decimal number, e.g. <code>0.444</code>.</em></p>

<h4>Part (c) &mdash; Debug Task (1 point, Evaluate)</h4>
<p>A student computed the self-inductance and obtained \(L_{\text{student}} = {@wrong_L_mH@}\;\text{mH}\).
Which <strong>wrong formula</strong> did they most likely use?</p>
<p>[[input:ans3]] [[validation:ans3]]</p>

<h4>Part (d) &mdash; Energy Scaling (1 point, Understand)</h4>
<p>If the current through the inductor is <strong>doubled</strong> (while \(L\) remains the same),
by what factor does the stored energy change?</p>
<p>[[input:ans4]] [[validation:ans4]]</p>
]]></text>
  </questiontext>

  <generalfeedback format="html">
    <text><![CDATA[
<h3>Worked Solution</h3>

<h4>Step 1: Cross-Sectional Area</h4>
<p>The solenoid has radius \(r = {@r_cm@}\;\text{cm} = {@r_val@}\;\text{m}\).</p>
<p>\[A = \pi r^2 = \pi \times ({@r_val@})^2 = {@A_disp@}\;\text{m}^2\]</p>

<h4>Step 2: Self-Inductance</h4>
<p>Using the formula for a solenoid:</p>
<p>\[L = \frac{\mu_0 N^2 A}{\ell} = \frac{4\pi \times 10^{-7} \times {@N_turns@}^2 \times {@A_disp@}}{@ell_val@}\]</p>
<p>\[L = {@L_val@}\;\text{H} = {@ta1@}\;\text{mH}\]</p>

<h4>Step 3: Energy Stored</h4>
<p>\[W = \tfrac{1}{2} L I^2 = \tfrac{1}{2} \times {@L_val@} \times ({@I_val@})^2 = {@W_val@}\;\text{J} = {@ta2@}\;\text{mJ}\]</p>

<h4>Part (c): Debug Task</h4>
<p>The student obtained \(L_{\text{student}} = {@wrong_L_mH@}\;\text{mH}\), which equals
\(\mu_0 N A / \ell\) (using \(N\) instead of \(N^2\)). The correct formula requires \(N^2\)
because self-inductance accounts for the flux <em>linkage</em> through all \(N\) turns.</p>

<h4>Part (d): Energy Scaling</h4>
<p>Since \(W = \tfrac{1}{2} L I^2\), doubling the current gives:</p>
<p>\[W' = \tfrac{1}{2} L (2I)^2 = \tfrac{1}{2} L \cdot 4I^2 = 4W\]</p>
<p>The stored energy <strong>quadruples</strong> because energy is proportional to \(I^2\).</p>

<hr />
<h4>Answer Summary</h4>
<table border="1" cellpadding="5" cellspacing="0">
  <tr><th>Part</th><th>Quantity</th><th>Correct Answer</th></tr>
  <tr><td>(a)</td><td>Self-inductance \(L\)</td><td>{@ta1@} mH</td></tr>
  <tr><td>(b)</td><td>Stored energy \(W\)</td><td>{@ta2@} mJ</td></tr>
  <tr><td>(c)</td><td>Wrong formula used</td><td>\(L = \mu_0 N A / \ell\) (used N instead of N&sup2;)</td></tr>
  <tr><td>(d)</td><td>Energy scaling factor</td><td>Energy quadruples (&times;4)</td></tr>
</table>
]]></text>
  </generalfeedback>

  <defaultgrade>4.0000000</defaultgrade>
  <penalty>0.0000000</penalty>
  <hidden>0</hidden>
  <idnumber/>

  <stackversion>
    <text>2024032101</text>
  </stackversion>

  <questionvariables>
    <text><![CDATA[
/* ===== Parameter set selection ===== */
set: rand(3)+1;

/* ===== Given values ===== */
N_turns: if set=1 then 200 else if set=2 then 500 else 300;
r_cm:    if set=1 then 2   else if set=2 then 1.5 else 3;
ell_cm:  if set=1 then 20  else if set=2 then 10  else 30;
I_val:   if set=1 then 1.5 else if set=2 then 0.8 else 2.0;

/* ===== Derived values ===== */
r_val:   r_cm/100;
ell_val: ell_cm/100;
A_val:   float(%pi*r_val^2);
mu0:     4*%pi/10^7;
L_val:   float(mu0*N_turns^2*A_val/ell_val);
W_val:   float(0.5*L_val*I_val^2);

/* Display-friendly values */
A_disp: float(A_val);

/* ===== Common errors for debug task ===== */
/* Wrong formula: L = mu0*N*A/ell (used N instead of N²) */
wrong_L:    float(mu0*N_turns*A_val/ell_val);
wrong_L_mH: float(wrong_L*1000);

/* ===== Teacher answers ===== */
ta1: float(L_val*1000);   /* L in mH */
ta2: float(W_val*1000);   /* W in mJ */

/* Part (c) MCQ: Debug task — which wrong formula was used? */
ta3: [[1, true,  "L = &mu;&#8320;NA/&ell; &mdash; they used N instead of N&sup2; in the formula"],
      [2, false, "L = &mu;&#8320;N&sup2;&ell;/A &mdash; they swapped the area and length"],
      [3, false, "L = &mu;&#8320;N&sup2;A&middot;&ell; &mdash; they multiplied by &ell; instead of dividing"],
      [4, false, "L = &mu;&#8320;A/(N&sup2;&ell;) &mdash; they put N&sup2; in the denominator"]];
ta3_correct: 1;

/* Part (d) MCQ: Energy scaling when current doubles */
ta4: [[1, false, "The energy doubles (&times;2)"],
      [2, true,  "The energy quadruples (&times;4) &mdash; energy is proportional to I&sup2;"],
      [3, false, "The energy stays the same &mdash; inductance hasn&rsquo;t changed"],
      [4, false, "The energy increases by a factor of 8"]];
ta4_correct: 2;
]]></text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[
<h4>Part (a) &mdash; Self-Inductance</h4>
<div>[[feedback:prt1]]</div>

<h4>Part (b) &mdash; Stored Energy</h4>
<div>[[feedback:prt2]]</div>

<h4>Part (c) &mdash; Debug Task</h4>
<div>[[feedback:prt3]]</div>

<h4>Part (d) &mdash; Energy Scaling</h4>
<div>[[feedback:prt4]]</div>
]]></text>
  </specificfeedback>

  <questionnote>
    <text>Set {@set@}: N={@N_turns@}, r={@r_cm@}cm, ℓ={@ell_cm@}cm, I={@I_val@}A | L={@ta1@} mH, W={@ta2@} mJ</text>
  </questionnote>

  <questiondescription format="html">
    <text><![CDATA[<p>Self-Inductance and Inductor Energy — Compute the self-inductance of a solenoid and the energy stored, debug a common formula error, and analyze energy scaling with current. Week 11 training exercise.</p>]]></text>
  </questiondescription>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>

  <decimals>.</decimals>
  <scientificnotation>*10</scientificnotation>
  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>j</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed/>

  <!-- ============================================================ -->
  <!-- INPUTS (4 total)                                              -->
  <!-- ============================================================ -->

  <!-- ans1: Part (a) — Self-inductance (numerical, in mH) -->
  <input>
    <name>ans1</name>
    <type>numerical</type>
    <tans>ta1</tans>
    <boxsize>10</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options/>
  </input>

  <!-- ans2: Part (b) — Stored energy (numerical, in mJ) -->
  <input>
    <name>ans2</name>
    <type>numerical</type>
    <tans>ta2</tans>
    <boxsize>10</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options/>
  </input>

  <!-- ans3: Part (c) — Debug task MCQ (radio) -->
  <input>
    <name>ans3</name>
    <type>radio</type>
    <tans>ta3</tans>
    <boxsize>3</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>1</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options/>
  </input>

  <!-- ans4: Part (d) — Energy scaling MCQ (radio) -->
  <input>
    <name>ans4</name>
    <type>radio</type>
    <tans>ta4</tans>
    <boxsize>3</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>1</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options/>
  </input>

  <!-- ============================================================ -->
  <!-- PRTs (4 total)                                                -->
  <!-- ============================================================ -->

  <!-- prt1: Part (a) — Self-inductance, tiered numerical -->
  <prt>
    <name>prt1</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text><![CDATA[
sa1_float: float(ans1);
ref1_float: float(ta1);
/* Order-of-magnitude check */
prt1_oom_ok: if (abs(ref1_float) > 1.0e-30 and abs(sa1_float) > 1.0e-30 and (abs(sa1_float/ref1_float - 10) < 1.5 or abs(sa1_float/ref1_float - 0.1) < 0.015)) then true else false;
]]></text>
    </feedbackvariables>
    <!-- Node 0: NumRelative 5% → full marks -->
    <node>
      <name>0</name>
      <description>NumRelative check L within 5%</description>
      <answertest>NumRelative</answertest>
      <sans>ans1</sans>
      <tans>ta1</tans>
      <testoptions>0.05</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(L = \mu_0 N^2 A / \ell = {@ta1@}\;\text{mH}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>1</falsenextnode>
      <falseanswernote>prt1-0-F</falseanswernote>
      <falsefeedback format="html">
        <text/>
      </falsefeedback>
    </node>
    <!-- Node 1: NumRelative 15% → partial credit -->
    <node>
      <name>1</name>
      <description>NumRelative check within 15%</description>
      <answertest>NumRelative</answertest>
      <sans>ans1</sans>
      <tans>ta1</tans>
      <testoptions>0.15</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.7000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-1-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Close! Your answer is within 15% of the expected value. Check your arithmetic or rounding. The exact answer is {@ta1@} mH.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>2</falsenextnode>
      <falseanswernote>prt1-1-F</falseanswernote>
      <falsefeedback format="html">
        <text/>
      </falsefeedback>
    </node>
    <!-- Node 2: Order-of-magnitude fallback -->
    <node>
      <name>2</name>
      <description>OoM fallback check</description>
      <answertest>AlgEquiv</answertest>
      <sans>prt1_oom_ok</sans>
      <tans>true</tans>
      <testoptions/>
      <quiet>1</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.3000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-2-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Your answer is off by about a factor of 10. Check your unit conversions (cm to m) or whether you converted H to mH correctly.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-2-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The self-inductance is \(L = \mu_0 N^2 A / \ell = {@ta1@}\;\text{mH}\). Make sure to convert cm to m for both the radius and the length, and use \(A = \pi r^2\).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt2: Part (b) — Stored energy, tiered numerical -->
  <prt>
    <name>prt2</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text><![CDATA[
sa2_float: float(ans2);
ref2_float: float(ta2);
/* Order-of-magnitude check */
prt2_oom_ok: if (abs(ref2_float) > 1.0e-30 and abs(sa2_float) > 1.0e-30 and (abs(sa2_float/ref2_float - 10) < 1.5 or abs(sa2_float/ref2_float - 0.1) < 0.015)) then true else false;
]]></text>
    </feedbackvariables>
    <!-- Node 0: NumRelative 5% → full marks -->
    <node>
      <name>0</name>
      <description>NumRelative check W within 5%</description>
      <answertest>NumRelative</answertest>
      <sans>ans2</sans>
      <tans>ta2</tans>
      <testoptions>0.05</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt2-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(W = \tfrac{1}{2} L I^2 = {@ta2@}\;\text{mJ}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>1</falsenextnode>
      <falseanswernote>prt2-0-F</falseanswernote>
      <falsefeedback format="html">
        <text/>
      </falsefeedback>
    </node>
    <!-- Node 1: NumRelative 15% → partial credit -->
    <node>
      <name>1</name>
      <description>NumRelative check within 15%</description>
      <answertest>NumRelative</answertest>
      <sans>ans2</sans>
      <tans>ta2</tans>
      <testoptions>0.15</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.7000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt2-1-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Close! Your answer is within 15% of the expected value. Check your arithmetic or rounding. The exact answer is {@ta2@} mJ.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>2</falsenextnode>
      <falseanswernote>prt2-1-F</falseanswernote>
      <falsefeedback format="html">
        <text/>
      </falsefeedback>
    </node>
    <!-- Node 2: Order-of-magnitude fallback -->
    <node>
      <name>2</name>
      <description>OoM fallback check</description>
      <answertest>AlgEquiv</answertest>
      <sans>prt2_oom_ok</sans>
      <tans>true</tans>
      <testoptions/>
      <quiet>1</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.3000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt2-2-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Your answer is off by about a factor of 10. Check your unit conversions or whether you converted J to mJ correctly.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt2-2-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The stored energy is \(W = \tfrac{1}{2} L I^2 = {@ta2@}\;\text{mJ}\). First compute \(L\) in henries, then use \(W = \tfrac{1}{2} L I^2\), and finally convert to mJ.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt3: Part (c) — Debug task MCQ -->
  <prt>
    <name>prt3</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check debug task MCQ answer</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans3</sans>
      <tans>ta3_correct</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt3-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! The student used \(L = \mu_0 N A / \ell\) instead of \(L = \mu_0 N^2 A / \ell\). They forgot to square \(N\). The ratio of correct to wrong is \(N = {@N_turns@}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt3-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. Compare the student's value \({@wrong_L_mH@}\;\text{mH}\) with the correct value \({@ta1@}\;\text{mH}\). The ratio is \({@ta1@}/{@wrong_L_mH@} = {@N_turns@} = N\). The student used \(N\) instead of \(N^2\) in the formula.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt4: Part (d) — Energy scaling MCQ -->
  <prt>
    <name>prt4</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check energy scaling MCQ answer</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans4</sans>
      <tans>ta4_correct</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt4-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! Since \(W = \tfrac{1}{2} L I^2\), doubling \(I\) gives \(W' = \tfrac{1}{2} L (2I)^2 = 4 \times \tfrac{1}{2} L I^2 = 4W\). The energy quadruples.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt4-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The energy stored in an inductor is \(W = \tfrac{1}{2} L I^2\). Since energy depends on \(I^2\), doubling the current multiplies the energy by \(2^2 = 4\). The energy <strong>quadruples</strong>.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- ============================================================ -->
  <!-- QUESTION TESTS                                                -->
  <!-- ============================================================ -->

  <!-- Test 1: All correct -->
  <qtest>
    <testcase>1</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 2: All wrong -->
  <qtest>
    <testcase>2</testcase>
    <testinput>
      <name>ans1</name>
      <value>0</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>0</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>2</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>1</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-2-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-2-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-0-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-0-F</expectedanswernote>
    </expected>
  </qtest>

  <!-- ============================================================ -->
  <!-- HINTS (3, progressive)                                        -->
  <!-- ============================================================ -->

  <hint format="html">
    <text><![CDATA[<p><strong>Hint 1 (Intuition):</strong> Self-inductance measures how much magnetic flux a coil creates per unit current through itself. Energy stored in an inductor is analogous to a &ldquo;magnetic spring&rdquo; &mdash; it stores energy in the magnetic field.</p>]]></text>
  </hint>
  <hint format="html">
    <text><![CDATA[<p><strong>Hint 2 (Formulas):</strong> For a solenoid: \(L = \mu_0 N^2 A / \ell\), where \(A = \pi r^2\). The stored energy is \(W = \tfrac{1}{2} L I^2\). Remember to convert cm to m for both the radius and the length!</p>]]></text>
  </hint>
  <hint format="html">
    <text><![CDATA[<p><strong>Hint 3 (Worked Step):</strong> First compute the cross-sectional area: \(A = \pi \times ({@r_val@})^2 = {@A_disp@}\;\text{m}^2\). Then \(L = 4\pi \times 10^{-7} \times {@N_turns@}^2 \times {@A_disp@} \,/\, {@ell_val@}\). Convert to mH. For the energy: \(W = \tfrac{1}{2} L I^2\), convert to mJ. For the debug task, compare the student&rsquo;s value with yours &mdash; what&rsquo;s different?</p>]]></text>
  </hint>

</question>
</quiz>