<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Q2 V1: Series RLC Natural Response, Underdamped -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q2V1: RLC Transient + Debug (Series RLC Natural Response, Underdamped)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q2V1: RLC Transient with Waveform Debug (Series RLC Natural Response, Underdamped)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 9 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.8.</em></p>
<hr />

<h4>Circuit Description</h4>
<p>A series RLC circuit is powered by a DC voltage source \(V_s\). The capacitor is fully charged in steady state. At \(t = 0\), a switch disconnects the source, leaving the series R-L-C loop closed for current to flow.</p>
<p>[DIAGRAM: Series RLC Natural Response, Underdamped &mdash; CircuiTikZ diagram to be added]</p>

<h4>Component Values</h4>
<ul>
  <li>\(R = 10\;\Omega\)</li>
  <li>\(L = 10\;\text{mH}\)</li>
  <li>\(C = 100\;\mu\text{F}\)</li>
  <li>Source: \(V_s = 20\;\text{V}\)</li>
</ul>

<h4>Initial Conditions</h4>
<p>Before switching: \(v_C(0^-) = V_s\), \(i_L(0^-) = 0\) (capacitor fully charged, no current in steady state).</p>

<hr />

<h4>Part (a) &mdash; Damping Parameters and Classification (2 points)</h4>
<p>Calculate the <strong>neper frequency</strong> (damping coefficient) \(\alpha\) and the <strong>resonant frequency</strong> \(\omega_0\).</p>
<p>\(\alpha = \) [[input:ans1]] \(\text{Np/s}\) &nbsp; [[validation:ans1]]</p>
<p><em>Enter a number, e.g. <code>1234</code></em></p>

<p>\(\omega_0 = \) [[input:ans2]] \(\text{rad/s}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>5678</code></em></p>

<p>Classify the damping regime: [[input:ans3]] &nbsp; [[validation:ans3]]</p>

<h4>Part (b) &mdash; Response Value (2 points)</h4>
<p>Calculate the current \(i(t_1)\) at \(t_1 = 1\;\text{ms}\).</p>
<p>\(i(t_1) = \) [[input:ans4]] &nbsp; [[validation:ans4]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (c) &mdash; Waveform Error Identification (2 points)</h4>
<p>A student attempted to plot the response and made errors. Their waveform is described below:</p>
<blockquote>A student plotted \(i(t)\) for \(t > 0\). Their plot shows:<br/>(i) The current starts at \(i(0^+) = 2\;\text{A}\) (but it should start at 0).<br/>(ii) The oscillation envelope decays with a time constant of 4 ms (but \(1/\alpha = 2\;\text{ms}\)).<br/>(iii) The oscillation frequency appears correct.<br/>(iv) The current approaches 0 as \(t \to \infty\) (correct for natural response).</blockquote>
<p><strong>Select ALL statements that identify actual errors in the student&rsquo;s waveform:</strong></p>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>

<h4>Part (d) &mdash; Parameter Sensitivity (3 points, manually graded)</h4>
<p>Suppose the resistance \(R\) is <strong>doubled</strong>.</p>
<ol>
  <li>How does this change \(\alpha\)?</li>
  <li>Does the damping regime change? If so, from what to what?</li>
  <li>Sketch the qualitative shape of the new response and explain whether energy dissipation is faster or slower.</li>
</ol>
<p>[[input:ans6]] &nbsp; [[validation:ans6]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V1 (Series RLC Natural Response, Underdamped)</h4>
<p><strong>Part (a):</strong></p>
<p>\(\\alpha = \\frac{R}{2L} = {@ta_alpha@}\;\text{Np/s}\)</p>
<p>\(\omega_0 = \frac{1}{\sqrt{LC}} = {@ta_omega0@}\;\text{rad/s}\)</p>
<p>Since \(\alpha < \omega_0\), the circuit is <strong>underdamped</strong>.</p>

<p><strong>Part (b):</strong> \(i(t_1) = {@ta_value@}\)</p>

<p><strong>Part (c):</strong> The errors are statements {@correct_errors@}.</p>

<p><strong>Part (d):</strong> When R doubles:</p>
<ul>
  <li>\(\alpha\) doubles (since \(\alpha = R/(2L)\))</li>
  <li>The damping regime may change depending on the new \(\alpha\) vs \(\omega_0\)</li>
  <li>More damping = faster energy dissipation, response decays quicker</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>9</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion>
    <text>2024032401</text>
  </stackversion>
  <questionvariables>
    <text><![CDATA[
/* V1: Series RLC Natural Response, Underdamped */
R_val: 10;
L_val: 0.01;
C_val: 0.0001;
Vs_val: 20;
t1: 0.001;

/* Damping parameters */
alpha_val: R_val/(2*L_val);
omega0_val: 1/sqrt(L_val*C_val);
omega_d: sqrt(omega0_val^2 - alpha_val^2);

/* Wrong alpha (other topology formula) for error detection */
alpha_wrong: 1/(2*R_val*C_val);

/* Regime classification MCQ */
regime_opts: random_permutation([[1, true, "Underdamped"], [2, false, "Overdamped"], [3, false, "Critically damped"]]);

/* Response at t1 */
val_at_t1: float((Vs_val/(omega_d*L_val))*exp(-alpha_val*t1)*sin(omega_d*t1));

/* Debug MCQ-MA */
debug_opts: random_permutation([[1, true, "The initial value i(0+) is wrong; it should be 0 for a fully charged capacitor"], [2, true, "The decay envelope uses the wrong time constant (too slow)"], [3, false, "The oscillation frequency does not match the damped frequency omega_d"], [4, false, "The final value should approach a nonzero DC level"], [5, false, "The waveform should not oscillate because the circuit is overdamped"]]);
correct_errors: {1,2};

/* Teacher answers */
ta_alpha: float(alpha_val);
ta_omega0: float(omega0_val);
ta_regime: regime_opts;
ta_value: val_at_t1;
ta_debug: correct_errors;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_alpha]]
[[feedback:prt_omega0]]
[[feedback:prt_regime]]
[[feedback:prt_value]]
[[feedback:prt_debug]]
]]></text>
  </specificfeedback>
  <questionnote>
    <text>{@ta_alpha@}, {@ta_omega0@}, {@ta_value@}, {@correct_errors@}</text>
  </questionnote>
  <questiondescription format="html">
    <text></text>
  </questiondescription>
  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_alpha</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>1234</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_omega0</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>5678</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_regime</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta_value</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>checkbox</type>
      <tans>ta_debug</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans6</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_alpha</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta_alpha</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\alpha\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_alpha-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>float(alpha_wrong)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You used the wrong formula. Check whether this is a series or parallel RLC circuit and use the correct \(\alpha\) formula.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_alpha-2-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect \(\alpha\). Review the damping coefficient formula for this circuit topology.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_omega0</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_omega0</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_omega0-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\omega_0\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_omega0-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(\omega_0 = 1/\sqrt{LC}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_regime</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_regime-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct classification.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_regime-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Compare \(\alpha\) with \(\omega_0\): if \(\alpha < \omega_0\), underdamped; if \(\alpha > \omega_0\), overdamped; if \(\alpha = \omega_0\), critically damped.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_value</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans4</sans>
        <tans>ta_value</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_value-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct response value.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_value-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Review the response formula for this damping regime and check your calculations.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_debug</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans5</sans>
        <tans>ta_debug</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_debug-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! You identified all waveform errors.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_debug-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all waveform errors were correctly identified.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q2 V2: Parallel RLC Step Response, Overdamped -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q2V2: RLC Transient + Debug (Parallel RLC Step Response, Overdamped)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q2V2: RLC Transient with Waveform Debug (Parallel RLC Step Response, Overdamped)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 9 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.8.</em></p>
<hr />

<h4>Circuit Description</h4>
<p>A parallel RLC circuit has zero initial conditions. At \(t = 0\), a DC current source \(I_s\) is switched on, delivering current to the parallel combination of R, L, and C.</p>
<p>[DIAGRAM: Parallel RLC Step Response, Overdamped &mdash; CircuiTikZ diagram to be added]</p>

<h4>Component Values</h4>
<ul>
  <li>\(R = 20\;\Omega\)</li>
  <li>\(L = 100\;\text{mH}\)</li>
  <li>\(C = 10\;\mu\text{F}\)</li>
  <li>Source: \(I_s = 100\;\text{mA}\)</li>
</ul>

<h4>Initial Conditions</h4>
<p>Before switching: \(v(0^-) = 0\), \(i_L(0^-) = 0\) (all zero initial conditions).</p>

<hr />

<h4>Part (a) &mdash; Damping Parameters and Classification (2 points)</h4>
<p>Calculate the <strong>neper frequency</strong> (damping coefficient) \(\alpha\) and the <strong>resonant frequency</strong> \(\omega_0\).</p>
<p>\(\alpha = \) [[input:ans1]] \(\text{Np/s}\) &nbsp; [[validation:ans1]]</p>
<p><em>Enter a number, e.g. <code>1234</code></em></p>

<p>\(\omega_0 = \) [[input:ans2]] \(\text{rad/s}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>5678</code></em></p>

<p>Classify the damping regime: [[input:ans3]] &nbsp; [[validation:ans3]]</p>

<h4>Part (b) &mdash; Response Value (2 points)</h4>
<p>Calculate the voltage \(v(t_1)\) across the parallel combination at \(t_1 = 0.5\;\text{ms}\).</p>
<p>\(v(t_1) = \) [[input:ans4]] &nbsp; [[validation:ans4]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (c) &mdash; Waveform Error Identification (2 points)</h4>
<p>A student attempted to plot the response and made errors. Their waveform is described below:</p>
<blockquote>A student plotted \(v(t)\) for \(t > 0\). Their plot shows:<br/>(i) The voltage starts at \(v(0^+) = 0\) (correct).<br/>(ii) The voltage oscillates with a damped sinusoid (but the circuit is overdamped!).<br/>(iii) The voltage approaches \(v(\infty) = I_s \cdot R = 2\;\text{V}\) (but at DC, the inductor is a short, so \(v(\infty) = 0\)).<br/>(iv) The voltage rises then falls back — double exponential shape.</blockquote>
<p><strong>Select ALL statements that identify actual errors in the student&rsquo;s waveform:</strong></p>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>

<h4>Part (d) &mdash; Parameter Sensitivity (3 points, manually graded)</h4>
<p>Suppose the resistance \(R\) is <strong>doubled</strong>.</p>
<ol>
  <li>How does this change \(\alpha\)?</li>
  <li>Does the damping regime change? If so, from what to what?</li>
  <li>Sketch the qualitative shape of the new response and explain whether energy dissipation is faster or slower.</li>
</ol>
<p>[[input:ans6]] &nbsp; [[validation:ans6]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V2 (Parallel RLC Step Response, Overdamped)</h4>
<p><strong>Part (a):</strong></p>
<p>\(\\alpha = \\frac{1}{2RC} = {@ta_alpha@}\;\text{Np/s}\)</p>
<p>\(\omega_0 = \frac{1}{\sqrt{LC}} = {@ta_omega0@}\;\text{rad/s}\)</p>
<p>Since \(\alpha > \omega_0\), the circuit is <strong>overdamped</strong>.</p>

<p><strong>Part (b):</strong> \(v(t_1) = {@ta_value@}\)</p>

<p><strong>Part (c):</strong> The errors are statements {@correct_errors@}.</p>

<p><strong>Part (d):</strong> When R doubles:</p>
<ul>
  <li>\(\alpha\) halves (since \(\alpha = 1/(2RC)\))</li>
  <li>The damping regime may change depending on the new \(\alpha\) vs \(\omega_0\)</li>
  <li>Less damping = slower energy dissipation, response oscillates more</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>9</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion>
    <text>2024032401</text>
  </stackversion>
  <questionvariables>
    <text><![CDATA[
/* V2: Parallel RLC Step Response, Overdamped */
R_val: 20;
L_val: 0.1;
C_val: 1e-05;
Is_val: 0.1;
t1: 0.0005;

/* Damping parameters */
alpha_val: 1/(2*R_val*C_val);
omega0_val: 1/sqrt(L_val*C_val);
omega_d: 0;  /* overdamped — no oscillation */

/* Wrong alpha (other topology formula) for error detection */
alpha_wrong: R_val/(2*L_val);

/* Regime classification MCQ */
regime_opts: random_permutation([[1, true, "Overdamped"], [2, false, "Underdamped"], [3, false, "Critically damped"]]);

/* Response at t1 */
disc: sqrt(alpha_val^2 - omega0_val^2);
s1: -alpha_val + disc;
s2: -alpha_val - disc;
A1: Is_val/(C_val*(s1 - s2));
A2: -A1;
val_at_t1: float(A1*exp(s1*t1) + A2*exp(s2*t1));

/* Debug MCQ-MA */
debug_opts: random_permutation([[1, false, "The initial value v(0+) = 0 is correct"], [2, true, "The waveform should NOT oscillate because the circuit is overdamped"], [3, true, "The final value v(infinity) should be 0, not I_s*R"], [4, false, "The damping coefficient alpha was computed with the wrong formula"], [5, false, "The waveform should start at a nonzero value"]]);
correct_errors: {2,3};

/* Teacher answers */
ta_alpha: float(alpha_val);
ta_omega0: float(omega0_val);
ta_regime: regime_opts;
ta_value: val_at_t1;
ta_debug: correct_errors;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_alpha]]
[[feedback:prt_omega0]]
[[feedback:prt_regime]]
[[feedback:prt_value]]
[[feedback:prt_debug]]
]]></text>
  </specificfeedback>
  <questionnote>
    <text>{@ta_alpha@}, {@ta_omega0@}, {@ta_value@}, {@correct_errors@}</text>
  </questionnote>
  <questiondescription format="html">
    <text></text>
  </questiondescription>
  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_alpha</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>1234</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_omega0</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>5678</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_regime</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta_value</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>checkbox</type>
      <tans>ta_debug</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans6</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_alpha</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta_alpha</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\alpha\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_alpha-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>float(alpha_wrong)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You used the wrong formula. Check whether this is a series or parallel RLC circuit and use the correct \(\alpha\) formula.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_alpha-2-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect \(\alpha\). Review the damping coefficient formula for this circuit topology.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_omega0</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_omega0</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_omega0-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\omega_0\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_omega0-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(\omega_0 = 1/\sqrt{LC}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_regime</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_regime-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct classification.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_regime-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Compare \(\alpha\) with \(\omega_0\): if \(\alpha < \omega_0\), underdamped; if \(\alpha > \omega_0\), overdamped; if \(\alpha = \omega_0\), critically damped.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_value</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans4</sans>
        <tans>ta_value</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_value-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct response value.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_value-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Review the response formula for this damping regime and check your calculations.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_debug</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans5</sans>
        <tans>ta_debug</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_debug-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! You identified all waveform errors.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_debug-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all waveform errors were correctly identified.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q2 V3: Series RLC Step Response, Critically Damped -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q2V3: RLC Transient + Debug (Series RLC Step Response, Critically Damped)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q2V3: RLC Transient with Waveform Debug (Series RLC Step Response, Critically Damped)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 9 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.8.</em></p>
<hr />

<h4>Circuit Description</h4>
<p>A series RLC circuit with zero initial conditions is connected to a DC voltage source \(V_s\) at \(t = 0\). The circuit response is the capacitor voltage \(v_C(t)\).</p>
<p>[DIAGRAM: Series RLC Step Response, Critically Damped &mdash; CircuiTikZ diagram to be added]</p>

<h4>Component Values</h4>
<ul>
  <li>\(R = 200\;\Omega\)</li>
  <li>\(L = 100\;\text{mH}\)</li>
  <li>\(C = 10\;\mu\text{F}\)</li>
  <li>Source: \(V_s = 10\;\text{V}\)</li>
</ul>

<h4>Initial Conditions</h4>
<p>Before switching: \(v_C(0^-) = 0\), \(i_L(0^-) = 0\) (zero initial conditions).</p>

<hr />

<h4>Part (a) &mdash; Damping Parameters and Classification (2 points)</h4>
<p>Calculate the <strong>neper frequency</strong> (damping coefficient) \(\alpha\) and the <strong>resonant frequency</strong> \(\omega_0\).</p>
<p>\(\alpha = \) [[input:ans1]] \(\text{Np/s}\) &nbsp; [[validation:ans1]]</p>
<p><em>Enter a number, e.g. <code>1234</code></em></p>

<p>\(\omega_0 = \) [[input:ans2]] \(\text{rad/s}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>5678</code></em></p>

<p>Classify the damping regime: [[input:ans3]] &nbsp; [[validation:ans3]]</p>

<h4>Part (b) &mdash; Response Value (2 points)</h4>
<p>Calculate the capacitor voltage \(v_C(t_1)\) at \(t_1 = 2\;\text{ms}\).</p>
<p>\(v_C(t_1) = \) [[input:ans4]] &nbsp; [[validation:ans4]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (c) &mdash; Waveform Error Identification (2 points)</h4>
<p>A student attempted to plot the response and made errors. Their waveform is described below:</p>
<blockquote>A student plotted \(v_C(t)\) for \(t > 0\). Their plot shows:<br/>(i) \(v_C(0^+) = V_s = 10\;\text{V}\) (but it should start at 0!).<br/>(ii) The voltage overshoots \(V_s\) before settling (but critically damped response does NOT overshoot!).<br/>(iii) The final value is \(v_C(\infty) = V_s = 10\;\text{V}\) (correct).<br/>(iv) The response reaches 63% of final value in about 1 ms.</blockquote>
<p><strong>Select ALL statements that identify actual errors in the student&rsquo;s waveform:</strong></p>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>

<h4>Part (d) &mdash; Parameter Sensitivity (3 points, manually graded)</h4>
<p>Suppose the resistance \(R\) is <strong>doubled</strong>.</p>
<ol>
  <li>How does this change \(\alpha\)?</li>
  <li>Does the damping regime change? If so, from what to what?</li>
  <li>Sketch the qualitative shape of the new response and explain whether energy dissipation is faster or slower.</li>
</ol>
<p>[[input:ans6]] &nbsp; [[validation:ans6]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V3 (Series RLC Step Response, Critically Damped)</h4>
<p><strong>Part (a):</strong></p>
<p>\(\\alpha = \\frac{R}{2L} = {@ta_alpha@}\;\text{Np/s}\)</p>
<p>\(\omega_0 = \frac{1}{\sqrt{LC}} = {@ta_omega0@}\;\text{rad/s}\)</p>
<p>Since \(\alpha = \omega_0\), the circuit is <strong>critically damped</strong>.</p>

<p><strong>Part (b):</strong> \(v_C(t_1) = {@ta_value@}\)</p>

<p><strong>Part (c):</strong> The errors are statements {@correct_errors@}.</p>

<p><strong>Part (d):</strong> When R doubles:</p>
<ul>
  <li>\(\alpha\) doubles (since \(\alpha = R/(2L)\))</li>
  <li>The damping regime may change depending on the new \(\alpha\) vs \(\omega_0\)</li>
  <li>More damping = faster energy dissipation, response decays quicker</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>9</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion>
    <text>2024032401</text>
  </stackversion>
  <questionvariables>
    <text><![CDATA[
/* V3: Series RLC Step Response, Critically Damped */
R_val: 200;
L_val: 0.1;
C_val: 1e-05;
Vs_val: 10;
t1: 0.002;

/* Damping parameters */
alpha_val: R_val/(2*L_val);
omega0_val: 1/sqrt(L_val*C_val);
omega_d: 0;  /* critically damped */

/* Wrong alpha (other topology formula) for error detection */
alpha_wrong: 1/(2*R_val*C_val);

/* Regime classification MCQ */
regime_opts: random_permutation([[1, true, "Critically damped"], [2, false, "Underdamped"], [3, false, "Overdamped"]]);

/* Response at t1 */
val_at_t1: float(Vs_val*(1 - (1 + alpha_val*t1)*exp(-alpha_val*t1)));

/* Debug MCQ-MA */
debug_opts: random_permutation([[1, true, "The initial value v_C(0+) should be 0, not V_s"], [2, true, "A critically damped response does NOT overshoot; the plot is wrong"], [3, false, "The final value v_C(infinity) = V_s is correct for a step response"], [4, false, "The time constant is wrong; it should be 1/alpha = 1 ms"], [5, false, "The circuit is actually underdamped, not critically damped"]]);
correct_errors: {1,2};

/* Teacher answers */
ta_alpha: float(alpha_val);
ta_omega0: float(omega0_val);
ta_regime: regime_opts;
ta_value: val_at_t1;
ta_debug: correct_errors;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_alpha]]
[[feedback:prt_omega0]]
[[feedback:prt_regime]]
[[feedback:prt_value]]
[[feedback:prt_debug]]
]]></text>
  </specificfeedback>
  <questionnote>
    <text>{@ta_alpha@}, {@ta_omega0@}, {@ta_value@}, {@correct_errors@}</text>
  </questionnote>
  <questiondescription format="html">
    <text></text>
  </questiondescription>
  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_alpha</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>1234</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_omega0</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>5678</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_regime</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta_value</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>checkbox</type>
      <tans>ta_debug</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans6</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_alpha</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta_alpha</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\alpha\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_alpha-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>float(alpha_wrong)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You used the wrong formula. Check whether this is a series or parallel RLC circuit and use the correct \(\alpha\) formula.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_alpha-2-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect \(\alpha\). Review the damping coefficient formula for this circuit topology.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_omega0</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_omega0</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_omega0-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\omega_0\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_omega0-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(\omega_0 = 1/\sqrt{LC}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_regime</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_regime-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct classification.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_regime-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Compare \(\alpha\) with \(\omega_0\): if \(\alpha < \omega_0\), underdamped; if \(\alpha > \omega_0\), overdamped; if \(\alpha = \omega_0\), critically damped.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_value</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans4</sans>
        <tans>ta_value</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_value-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct response value.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_value-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Review the response formula for this damping regime and check your calculations.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_debug</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans5</sans>
        <tans>ta_debug</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_debug-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! You identified all waveform errors.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_debug-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all waveform errors were correctly identified.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q2 V4: Parallel RLC Natural Response, Underdamped -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q2V4: RLC Transient + Debug (Parallel RLC Natural Response, Underdamped)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q2V4: RLC Transient with Waveform Debug (Parallel RLC Natural Response, Underdamped)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 9 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.8.</em></p>
<hr />

<h4>Circuit Description</h4>
<p>A parallel RLC circuit has a capacitor pre-charged to \(V_0\). At \(t = 0\), a switch closes, connecting the capacitor to R and L in parallel. The voltage \(v(t)\) across the combination is the response.</p>
<p>[DIAGRAM: Parallel RLC Natural Response, Underdamped &mdash; CircuiTikZ diagram to be added]</p>

<h4>Component Values</h4>
<ul>
  <li>\(R = 100\;\Omega\)</li>
  <li>\(L = 50\;\text{mH}\)</li>
  <li>\(C = 5\;\mu\text{F}\)</li>
  <li>Source: \(V_0 = 15\;\text{V}\)</li>
</ul>

<h4>Initial Conditions</h4>
<p>Before switching: \(v(0^-) = V_0 = 15\;\text{V}\), \(i_L(0^-) = 0\).</p>

<hr />

<h4>Part (a) &mdash; Damping Parameters and Classification (2 points)</h4>
<p>Calculate the <strong>neper frequency</strong> (damping coefficient) \(\alpha\) and the <strong>resonant frequency</strong> \(\omega_0\).</p>
<p>\(\alpha = \) [[input:ans1]] \(\text{Np/s}\) &nbsp; [[validation:ans1]]</p>
<p><em>Enter a number, e.g. <code>1234</code></em></p>

<p>\(\omega_0 = \) [[input:ans2]] \(\text{rad/s}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>5678</code></em></p>

<p>Classify the damping regime: [[input:ans3]] &nbsp; [[validation:ans3]]</p>

<h4>Part (b) &mdash; Response Value (2 points)</h4>
<p>Calculate the voltage \(v(t_1)\) at \(t_1 = 0.5\;\text{ms}\).</p>
<p>\(v(t_1) = \) [[input:ans4]] &nbsp; [[validation:ans4]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (c) &mdash; Waveform Error Identification (2 points)</h4>
<p>A student attempted to plot the response and made errors. Their waveform is described below:</p>
<blockquote>A student plotted \(v(t)\) for \(t > 0\). Their plot shows:<br/>(i) The voltage starts at \(v(0^+) = V_0 = 15\;\text{V}\) (correct).<br/>(ii) The envelope decays with time constant \(1/\omega_0\) instead of \(1/\alpha\).<br/>(iii) The oscillation frequency is \(\omega_0\) instead of the damped frequency \(\omega_d = \sqrt{\omega_0^2 - \alpha^2}\).<br/>(iv) The voltage approaches 0 as \(t \to \infty\) (correct for natural response).</blockquote>
<p><strong>Select ALL statements that identify actual errors in the student&rsquo;s waveform:</strong></p>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>

<h4>Part (d) &mdash; Parameter Sensitivity (3 points, manually graded)</h4>
<p>Suppose the resistance \(R\) is <strong>doubled</strong>.</p>
<ol>
  <li>How does this change \(\alpha\)?</li>
  <li>Does the damping regime change? If so, from what to what?</li>
  <li>Sketch the qualitative shape of the new response and explain whether energy dissipation is faster or slower.</li>
</ol>
<p>[[input:ans6]] &nbsp; [[validation:ans6]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V4 (Parallel RLC Natural Response, Underdamped)</h4>
<p><strong>Part (a):</strong></p>
<p>\(\\alpha = \\frac{1}{2RC} = {@ta_alpha@}\;\text{Np/s}\)</p>
<p>\(\omega_0 = \frac{1}{\sqrt{LC}} = {@ta_omega0@}\;\text{rad/s}\)</p>
<p>Since \(\alpha < \omega_0\), the circuit is <strong>underdamped</strong>.</p>

<p><strong>Part (b):</strong> \(v(t_1) = {@ta_value@}\)</p>

<p><strong>Part (c):</strong> The errors are statements {@correct_errors@}.</p>

<p><strong>Part (d):</strong> When R doubles:</p>
<ul>
  <li>\(\alpha\) halves (since \(\alpha = 1/(2RC)\))</li>
  <li>The damping regime may change depending on the new \(\alpha\) vs \(\omega_0\)</li>
  <li>Less damping = slower energy dissipation, response oscillates more</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>9</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion>
    <text>2024032401</text>
  </stackversion>
  <questionvariables>
    <text><![CDATA[
/* V4: Parallel RLC Natural Response, Underdamped */
R_val: 100;
L_val: 0.05;
C_val: 5e-06;
V0_val: 15;
t1: 0.0005;

/* Damping parameters */
alpha_val: 1/(2*R_val*C_val);
omega0_val: 1/sqrt(L_val*C_val);
omega_d: sqrt(omega0_val^2 - alpha_val^2);

/* Wrong alpha (other topology formula) for error detection */
alpha_wrong: R_val/(2*L_val);

/* Regime classification MCQ */
regime_opts: random_permutation([[1, true, "Underdamped"], [2, false, "Overdamped"], [3, false, "Critically damped"]]);

/* Response at t1 */
omega_d: sqrt(omega0_val^2 - alpha_val^2);
val_at_t1: float(V0_val*exp(-alpha_val*t1)*(cos(omega_d*t1) + (alpha_val/omega_d)*sin(omega_d*t1)));

/* Debug MCQ-MA */
debug_opts: random_permutation([[1, false, "The initial value v(0+) = V_0 is correct"], [2, true, "The decay envelope uses omega_0 instead of alpha for the time constant"], [3, true, "The oscillation frequency should be omega_d, not omega_0"], [4, false, "The final value should be nonzero"], [5, false, "The waveform should be a pure exponential (no oscillation)"]]);
correct_errors: {2,3};

/* Teacher answers */
ta_alpha: float(alpha_val);
ta_omega0: float(omega0_val);
ta_regime: regime_opts;
ta_value: val_at_t1;
ta_debug: correct_errors;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_alpha]]
[[feedback:prt_omega0]]
[[feedback:prt_regime]]
[[feedback:prt_value]]
[[feedback:prt_debug]]
]]></text>
  </specificfeedback>
  <questionnote>
    <text>{@ta_alpha@}, {@ta_omega0@}, {@ta_value@}, {@correct_errors@}</text>
  </questionnote>
  <questiondescription format="html">
    <text></text>
  </questiondescription>
  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_alpha</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>1234</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_omega0</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>5678</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_regime</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta_value</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>checkbox</type>
      <tans>ta_debug</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans6</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_alpha</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta_alpha</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\alpha\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_alpha-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>float(alpha_wrong)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_alpha-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You used the wrong formula. Check whether this is a series or parallel RLC circuit and use the correct \(\alpha\) formula.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_alpha-2-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect \(\alpha\). Review the damping coefficient formula for this circuit topology.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_omega0</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_omega0</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_omega0-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct \(\omega_0\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_omega0-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(\omega_0 = 1/\sqrt{LC}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_regime</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_regime-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct classification.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_regime-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Compare \(\alpha\) with \(\omega_0\): if \(\alpha < \omega_0\), underdamped; if \(\alpha > \omega_0\), overdamped; if \(\alpha = \omega_0\), critically damped.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_value</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans4</sans>
        <tans>ta_value</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_value-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct response value.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_value-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Review the response formula for this damping regime and check your calculations.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_debug</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans5</sans>
        <tans>ta_debug</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_debug-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! You identified all waveform errors.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_debug-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all waveform errors were correctly identified.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Companion Essay: Handwritten Upload for Q2 -->
<!-- ============================================================ -->
<question type="essay">
  <name>
    <text>Midterm 2 - Q2: Upload Handwritten Work (RLC Transient)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Upload Your Handwritten Work for Q2</h3>

<p>Upload clear photos or PDF scans of your <strong>handwritten work</strong> for Question 2 (RLC Transient with Waveform Debug).</p>

<hr />

<h4>Your uploads must include:</h4>
<ul>
  <li><strong>Part (a):</strong>
    <ul>
      <li>\(\alpha\) calculation with the correct formula for the circuit topology</li>
      <li>\(\omega_0\) calculation</li>
      <li>Classification logic: comparison of \(\alpha\) vs \(\omega_0\)</li>
    </ul>
  </li>
  <li><strong>Part (b):</strong>
    <ul>
      <li>Response formula setup (which form: overdamped, underdamped, or critically damped)</li>
      <li>Coefficient calculation from initial conditions</li>
      <li>Evaluation at \(t_1\)</li>
    </ul>
  </li>
  <li><strong>Part (d):</strong>
    <ul>
      <li>New \(\alpha\) calculation with doubled R</li>
      <li>Qualitative waveform sketch showing the change in behavior</li>
      <li>Energy dissipation reasoning</li>
    </ul>
  </li>
</ul>

<h4>Upload Instructions:</h4>
<ol>
  <li>Write your <strong>name and student ID</strong> on every page.</li>
  <li>Use a <strong>dark pen</strong> (blue or black) on white paper.</li>
  <li>Ensure photos are <strong>legible</strong>.</li>
  <li>Up to <strong>2 files</strong>. Accepted: <strong>PDF, JPG, JPEG, PNG</strong>.</li>
</ol>

<p><em>Worth 0 points. Supports grading verification.</em></p>
<p><strong>Teacher notice:</strong> Set &ldquo;Allow attachments&rdquo; to 2+ in quiz settings.</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[<p>Thank you for uploading your handwritten work for Q2.</p>]]></text>
  </generalfeedback>
  <defaultgrade>0</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <responseformat>noinline</responseformat>
  <responserequired>0</responserequired>
  <responsefieldlines>5</responsefieldlines>
  <attachments>2</attachments>
  <attachmentsrequired>1</attachmentsrequired>
  <filetypeslist>.pdf,.jpg,.jpeg,.png</filetypeslist>
  <graderinfo format="html">
    <text><![CDATA[
<h4>Instructor Notes &mdash; Q2 Handwritten Work Review</h4>
<ul>
  <li>Check correct \(\alpha\) formula for the topology (series vs parallel)</li>
  <li>Verify response form matches the damping regime</li>
  <li>For Part (d): look for correct reasoning about how R affects \(\alpha\) and the regime</li>
  <li><strong>Red flags:</strong> Answers without supporting calculations, inconsistent handwriting</li>
</ul>
]]></text>
  </graderinfo>
  <responsetemplate format="html">
    <text></text>
  </responsetemplate>
</question>

</quiz>