<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Q3 V1: E-core, Aiding Dots, Resistive Load -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q3V1: Coupled Circuits + Transformer (E-core, Aiding Dots, Resistive Load)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q3V1: Coupled Circuits + Transformer Design (E-core, Aiding Dots, Resistive Load)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 10 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.9 Problems 9.6&ndash;9.24.</em></p>
<hr />

<h4>Scenario</h4>
<p>Two magnetically coupled coils have self-inductances \(L_1 = 50\;\text{mH}\) and
\(L_2 = 20\;\text{mH}\), with coupling coefficient \(k = {@k_val@}\).</p>

<p>Both coils are wound in the <strong>same sense</strong> (clockwise when viewed from the top). Current \(i_1\) enters the top terminal of Coil 1, and \(i_2\) enters the top terminal of Coil 2.</p>

<p>[DIAGRAM: Physical winding diagram for E-core, Aiding Dots, Resistive Load &mdash; to be added]</p>

<p>The coils carry DC currents \(i_1 = 3\;\text{A}\) and \(i_2 = 2\;\text{A}\).</p>

<hr />

<h4>Part (a) &mdash; Dot Configuration (1 point)</h4>
<p>Based on the winding description above, select the correct dot placement:</p>
<p>[[input:ans1]] &nbsp; [[validation:ans1]]</p>
<p><em>(This part will be upgraded to drag-and-drop when diagrams are available.)</em></p>

<h4>Part (b) &mdash; Mutual Inductance (2 points)</h4>
<p>Calculate the mutual inductance \(M\).</p>
<p>\(M = \) [[input:ans2]] \(\text{H}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>0.0456</code></em></p>

<h4>Part (c) &mdash; Stored Energy (2 points)</h4>
<p>Calculate the total magnetic energy stored in the coupled system.</p>
<p>Recall: \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\)</p>
<p>The sign of the mutual term depends on the dot convention and current directions.</p>
<p>\(W = \) [[input:ans3]] \(\text{J}\) &nbsp; [[validation:ans3]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (d) &mdash; Energy Optimization (1 point)</h4>
<p><strong>Select ALL changes that would increase the total stored energy:</strong></p>
<p>[[input:ans4]] &nbsp; [[validation:ans4]]</p>

<h4>Part (e) &mdash; Ideal Transformer Comparison (4 points, manually graded)</h4>
<p>An engineer proposes modeling these coupled coils as an <strong>ideal transformer</strong> with
turns ratio \(a = \sqrt{L_1/L_2} = {@float(a_val)@}\).</p>
<ol>
  <li>Under what conditions is the ideal transformer approximation valid?</li>
  <li>What physical effects are lost in the ideal transformer model?</li>
  <li>For the given currents, what is the relationship between \(i_1\) and \(i_2\)
      under the ideal transformer assumption? Does this match the given values?</li>
</ol>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V1 (E-core, Aiding Dots, Resistive Load)</h4>
<p><strong>Part (a):</strong> The dot configuration is aiding.</p>
<p><strong>Part (b):</strong> \(M = k\sqrt{L_1 L_2} = {@M_num@}\;\text{H}\)</p>
<p><strong>Part (c):</strong></p>
<p>\(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 + M i_1 i_2 = {@W_num@}\;\text{J}\)</p>
<p>The sign is <strong>+</strong> because the dot configuration is <strong>aiding</strong>.</p>
<p><strong>Part (d):</strong> The correct selections are {@correct_mcq@}.</p>
<p><strong>Part (e):</strong></p>
<ul>
  <li>The ideal transformer requires \(k \to 1\) (perfect coupling) and infinite inductances with finite ratio</li>
  <li>Lost effects: leakage flux, core losses, winding resistance, finite inductance (DC magnetization)</li>
  <li>Ideal transformer: \(i_1/i_2 = 1/a = {@float(1/a_val)@}\). The given currents \(i_1 = 3\) A and \(i_2 = 2\) A may not satisfy this ratio, showing the coils are not ideal.</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>10</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2024032401</text></stackversion>
  <questionvariables>
    <text><![CDATA[
/* V1: E-core, Aiding Dots, Resistive Load */
L1_val: 0.05;
L2_val: 0.02;
k_val: 0.6;
i1_val: 3;
i2_val: 2;

/* Mutual inductance */
M_val: k_val*sqrt(L1_val*L2_val);

/* Stored energy: W = 1/2*L1*i1^2 + 1/2*L2*i2^2 + M*i1*i2 */
W_val: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 + M_val*i1_val*i2_val;

/* Error values */
W_wrong_sign: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 - M_val*i1_val*i2_val;
W_forgot_M: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2;

/* T-equivalent elements */
La_val: L1_val - M_val;
Lb_val: L2_val - M_val;

/* Turns ratio estimate */
a_val: sqrt(L1_val/L2_val);

/* Numerical versions */
M_num: float(M_val);
W_num: float(W_val);

/* Part (a) dot config MCQ */
dot_opts: random_permutation([[1, true, "Dots on top terminals of both coils (same-sense winding)"], [2, false, "Dot on top of Coil 1, dot on bottom of Coil 2"], [3, false, "Dots on bottom terminals of both coils"], [4, false, "Dot on bottom of Coil 1, dot on top of Coil 2"]]);

/* Part (d) MCQ-MA */
mcq_opts: random_permutation([[1, true, "Increase the coupling coefficient k"], [2, false, "Reverse the direction of i_2"], [3, true, "Increase L_1"], [4, false, "Decrease both i_1 and i_2 by half"], [5, false, "Replace aiding dots with opposing dots"]]);
correct_mcq: {1,3};

/* Teacher answers */
ta_dot: dot_opts;
ta_M: float(M_val);
ta_W: float(W_val);
ta_mcq: correct_mcq;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_dot]]
[[feedback:prt_M]]
[[feedback:prt_W]]
[[feedback:prt_mcq]]
]]></text>
  </specificfeedback>
  <questionnote><text>{@M_num@}, {@W_num@}, {@correct_mcq@}</text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <prtcorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text></prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_dot</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_M</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.0456</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_W</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>checkbox</type>
      <tans>ta_mcq</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_dot</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_dot-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct dot placement.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_dot-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. The dot marks the terminal where current entering produces flux in the reference direction. Check the winding sense.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_M</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_M</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_M-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct mutual inductance.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_M-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(M = k\sqrt{L_1 L_2}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_W</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>ta_W</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct stored energy.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_W-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_wrong_sign)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Wrong sign on the mutual term! Check whether the configuration is aiding or opposing. Aiding: +M, Opposing: -M.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_W-2-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_forgot_M)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-3-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You computed only the self-energy terms and forgot the mutual energy term \(\pm M i_1 i_2\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_W-3-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_mcq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans4</sans>
        <tans>ta_mcq</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mcq-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! All energy-increasing changes identified.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mcq-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all correct changes identified. Consider how each parameter affects the energy formula.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q3 V2: Toroid, Opposing Dots, RL Load -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q3V2: Coupled Circuits + Transformer (Toroid, Opposing Dots, RL Load)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q3V2: Coupled Circuits + Transformer Design (Toroid, Opposing Dots, RL Load)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 10 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.9 Problems 9.6&ndash;9.24.</em></p>
<hr />

<h4>Scenario</h4>
<p>Two magnetically coupled coils have self-inductances \(L_1 = 40\;\text{mH}\) and
\(L_2 = 10\;\text{mH}\), with coupling coefficient \(k = {@k_val@}\).</p>

<p>The coils are wound in <strong>opposite senses</strong> on a toroidal core. Current \(i_1\) enters the dotted terminal of Coil 1, but \(i_2\) enters the <strong>undotted</strong> terminal of Coil 2.</p>

<p>[DIAGRAM: Physical winding diagram for Toroid, Opposing Dots, RL Load &mdash; to be added]</p>

<p>The coils carry DC currents \(i_1 = 4\;\text{A}\) and \(i_2 = 3\;\text{A}\).</p>

<hr />

<h4>Part (a) &mdash; Dot Configuration (1 point)</h4>
<p>Based on the winding description above, select the correct dot placement:</p>
<p>[[input:ans1]] &nbsp; [[validation:ans1]]</p>
<p><em>(This part will be upgraded to drag-and-drop when diagrams are available.)</em></p>

<h4>Part (b) &mdash; Mutual Inductance (2 points)</h4>
<p>Calculate the mutual inductance \(M\).</p>
<p>\(M = \) [[input:ans2]] \(\text{H}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>0.0456</code></em></p>

<h4>Part (c) &mdash; Stored Energy (2 points)</h4>
<p>Calculate the total magnetic energy stored in the coupled system.</p>
<p>Recall: \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\)</p>
<p>The sign of the mutual term depends on the dot convention and current directions.</p>
<p>\(W = \) [[input:ans3]] \(\text{J}\) &nbsp; [[validation:ans3]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (d) &mdash; Energy Optimization (1 point)</h4>
<p><strong>Select ALL changes that would increase the total stored energy:</strong></p>
<p>[[input:ans4]] &nbsp; [[validation:ans4]]</p>

<h4>Part (e) &mdash; Ideal Transformer Comparison (4 points, manually graded)</h4>
<p>An engineer proposes modeling these coupled coils as an <strong>ideal transformer</strong> with
turns ratio \(a = \sqrt{L_1/L_2} = {@float(a_val)@}\).</p>
<ol>
  <li>Under what conditions is the ideal transformer approximation valid?</li>
  <li>What physical effects are lost in the ideal transformer model?</li>
  <li>For the given currents, what is the relationship between \(i_1\) and \(i_2\)
      under the ideal transformer assumption? Does this match the given values?</li>
</ol>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V2 (Toroid, Opposing Dots, RL Load)</h4>
<p><strong>Part (a):</strong> The dot configuration is opposing.</p>
<p><strong>Part (b):</strong> \(M = k\sqrt{L_1 L_2} = {@M_num@}\;\text{H}\)</p>
<p><strong>Part (c):</strong></p>
<p>\(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 - M i_1 i_2 = {@W_num@}\;\text{J}\)</p>
<p>The sign is <strong>-</strong> because the dot configuration is <strong>opposing</strong>.</p>
<p><strong>Part (d):</strong> The correct selections are {@correct_mcq@}.</p>
<p><strong>Part (e):</strong></p>
<ul>
  <li>The ideal transformer requires \(k \to 1\) (perfect coupling) and infinite inductances with finite ratio</li>
  <li>Lost effects: leakage flux, core losses, winding resistance, finite inductance (DC magnetization)</li>
  <li>Ideal transformer: \(i_1/i_2 = 1/a = {@float(1/a_val)@}\). The given currents \(i_1 = 4\) A and \(i_2 = 3\) A may not satisfy this ratio, showing the coils are not ideal.</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>10</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2024032401</text></stackversion>
  <questionvariables>
    <text><![CDATA[
/* V2: Toroid, Opposing Dots, RL Load */
L1_val: 0.04;
L2_val: 0.01;
k_val: 0.5;
i1_val: 4;
i2_val: 3;

/* Mutual inductance */
M_val: k_val*sqrt(L1_val*L2_val);

/* Stored energy: W = 1/2*L1*i1^2 + 1/2*L2*i2^2 - M*i1*i2 */
W_val: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 - M_val*i1_val*i2_val;

/* Error values */
W_wrong_sign: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 + M_val*i1_val*i2_val;
W_forgot_M: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2;

/* T-equivalent elements */
La_val: L1_val + M_val;
Lb_val: L2_val + M_val;

/* Turns ratio estimate */
a_val: sqrt(L1_val/L2_val);

/* Numerical versions */
M_num: float(M_val);
W_num: float(W_val);

/* Part (a) dot config MCQ */
dot_opts: random_permutation([[1, false, "Dots on top terminals of both coils"], [2, true, "Dot on top of Coil 1, dot on bottom of Coil 2 (opposite sense)"], [3, false, "Dots on bottom terminals of both coils"], [4, false, "Dot on bottom of Coil 1, dot on top of Coil 2"]]);

/* Part (d) MCQ-MA */
mcq_opts: random_permutation([[1, false, "Increase the coupling coefficient k"], [2, true, "Increase i_1"], [3, false, "Decrease L_2"], [4, true, "Increase L_1"], [5, false, "Reverse the direction of i_2 (makes both enter dotted terminals)"]]);
correct_mcq: {2,4};

/* Teacher answers */
ta_dot: dot_opts;
ta_M: float(M_val);
ta_W: float(W_val);
ta_mcq: correct_mcq;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_dot]]
[[feedback:prt_M]]
[[feedback:prt_W]]
[[feedback:prt_mcq]]
]]></text>
  </specificfeedback>
  <questionnote><text>{@M_num@}, {@W_num@}, {@correct_mcq@}</text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <prtcorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text></prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_dot</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_M</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.0456</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_W</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>checkbox</type>
      <tans>ta_mcq</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_dot</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_dot-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct dot placement.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_dot-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. The dot marks the terminal where current entering produces flux in the reference direction. Check the winding sense.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_M</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_M</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_M-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct mutual inductance.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_M-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(M = k\sqrt{L_1 L_2}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_W</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>ta_W</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct stored energy.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_W-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_wrong_sign)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Wrong sign on the mutual term! Check whether the configuration is aiding or opposing. Aiding: +M, Opposing: -M.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_W-2-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_forgot_M)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-3-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You computed only the self-energy terms and forgot the mutual energy term \(\pm M i_1 i_2\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_W-3-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_mcq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans4</sans>
        <tans>ta_mcq</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mcq-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! All energy-increasing changes identified.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mcq-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all correct changes identified. Consider how each parameter affects the energy formula.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q3 V3: C-core, Aiding Dots, Symmetric -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q3V3: Coupled Circuits + Transformer (C-core, Aiding Dots, Symmetric)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q3V3: Coupled Circuits + Transformer Design (C-core, Aiding Dots, Symmetric)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 10 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.9 Problems 9.6&ndash;9.24.</em></p>
<hr />

<h4>Scenario</h4>
<p>Two magnetically coupled coils have self-inductances \(L_1 = 30\;\text{mH}\) and
\(L_2 = 30\;\text{mH}\), with coupling coefficient \(k = {@k_val@}\).</p>

<p>Both coils have the <strong>same number of turns</strong> and are wound in the <strong>same direction</strong> on a C-core. Both currents enter their dotted terminals.</p>

<p>[DIAGRAM: Physical winding diagram for C-core, Aiding Dots, Symmetric &mdash; to be added]</p>

<p>The coils carry DC currents \(i_1 = 2\;\text{A}\) and \(i_2 = 2\;\text{A}\).</p>

<hr />

<h4>Part (a) &mdash; Dot Configuration (1 point)</h4>
<p>Based on the winding description above, select the correct dot placement:</p>
<p>[[input:ans1]] &nbsp; [[validation:ans1]]</p>
<p><em>(This part will be upgraded to drag-and-drop when diagrams are available.)</em></p>

<h4>Part (b) &mdash; Mutual Inductance (2 points)</h4>
<p>Calculate the mutual inductance \(M\).</p>
<p>\(M = \) [[input:ans2]] \(\text{H}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>0.0456</code></em></p>

<h4>Part (c) &mdash; Stored Energy (2 points)</h4>
<p>Calculate the total magnetic energy stored in the coupled system.</p>
<p>Recall: \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\)</p>
<p>The sign of the mutual term depends on the dot convention and current directions.</p>
<p>\(W = \) [[input:ans3]] \(\text{J}\) &nbsp; [[validation:ans3]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (d) &mdash; Energy Optimization (1 point)</h4>
<p><strong>Select ALL changes that would increase the total stored energy:</strong></p>
<p>[[input:ans4]] &nbsp; [[validation:ans4]]</p>

<h4>Part (e) &mdash; Ideal Transformer Comparison (4 points, manually graded)</h4>
<p>An engineer proposes modeling these coupled coils as an <strong>ideal transformer</strong> with
turns ratio \(a = \sqrt{L_1/L_2} = {@float(a_val)@}\).</p>
<ol>
  <li>Under what conditions is the ideal transformer approximation valid?</li>
  <li>What physical effects are lost in the ideal transformer model?</li>
  <li>For the given currents, what is the relationship between \(i_1\) and \(i_2\)
      under the ideal transformer assumption? Does this match the given values?</li>
</ol>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V3 (C-core, Aiding Dots, Symmetric)</h4>
<p><strong>Part (a):</strong> The dot configuration is aiding.</p>
<p><strong>Part (b):</strong> \(M = k\sqrt{L_1 L_2} = {@M_num@}\;\text{H}\)</p>
<p><strong>Part (c):</strong></p>
<p>\(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 + M i_1 i_2 = {@W_num@}\;\text{J}\)</p>
<p>The sign is <strong>+</strong> because the dot configuration is <strong>aiding</strong>.</p>
<p><strong>Part (d):</strong> The correct selections are {@correct_mcq@}.</p>
<p><strong>Part (e):</strong></p>
<ul>
  <li>The ideal transformer requires \(k \to 1\) (perfect coupling) and infinite inductances with finite ratio</li>
  <li>Lost effects: leakage flux, core losses, winding resistance, finite inductance (DC magnetization)</li>
  <li>Ideal transformer: \(i_1/i_2 = 1/a = {@float(1/a_val)@}\). The given currents \(i_1 = 2\) A and \(i_2 = 2\) A may not satisfy this ratio, showing the coils are not ideal.</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>10</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2024032401</text></stackversion>
  <questionvariables>
    <text><![CDATA[
/* V3: C-core, Aiding Dots, Symmetric */
L1_val: 0.03;
L2_val: 0.03;
k_val: 0.7;
i1_val: 2;
i2_val: 2;

/* Mutual inductance */
M_val: k_val*sqrt(L1_val*L2_val);

/* Stored energy: W = 1/2*L1*i1^2 + 1/2*L2*i2^2 + M*i1*i2 */
W_val: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 + M_val*i1_val*i2_val;

/* Error values */
W_wrong_sign: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 - M_val*i1_val*i2_val;
W_forgot_M: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2;

/* T-equivalent elements */
La_val: L1_val - M_val;
Lb_val: L2_val - M_val;

/* Turns ratio estimate */
a_val: sqrt(L1_val/L2_val);

/* Numerical versions */
M_num: float(M_val);
W_num: float(W_val);

/* Part (a) dot config MCQ */
dot_opts: random_permutation([[1, true, "Dots on left terminals of both coils (same winding direction)"], [2, false, "Dot on left of Coil 1, dot on right of Coil 2"], [3, false, "Dots on right terminals of both coils"], [4, false, "No dots needed since L_1 = L_2"]]);

/* Part (d) MCQ-MA */
mcq_opts: random_permutation([[1, true, "Increase k toward 1.0 (tighter coupling)"], [2, true, "Increase both i_1 and i_2 equally"], [3, false, "Reduce L_1 while keeping L_2 the same"], [4, false, "Switch from aiding to opposing dot convention"], [5, false, "Decrease the frequency of the source"]]);
correct_mcq: {1,2};

/* Teacher answers */
ta_dot: dot_opts;
ta_M: float(M_val);
ta_W: float(W_val);
ta_mcq: correct_mcq;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_dot]]
[[feedback:prt_M]]
[[feedback:prt_W]]
[[feedback:prt_mcq]]
]]></text>
  </specificfeedback>
  <questionnote><text>{@M_num@}, {@W_num@}, {@correct_mcq@}</text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <prtcorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text></prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_dot</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_M</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.0456</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_W</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>checkbox</type>
      <tans>ta_mcq</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_dot</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_dot-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct dot placement.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_dot-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. The dot marks the terminal where current entering produces flux in the reference direction. Check the winding sense.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_M</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_M</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_M-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct mutual inductance.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_M-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(M = k\sqrt{L_1 L_2}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_W</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>ta_W</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct stored energy.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_W-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_wrong_sign)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Wrong sign on the mutual term! Check whether the configuration is aiding or opposing. Aiding: +M, Opposing: -M.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_W-2-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_forgot_M)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-3-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You computed only the self-energy terms and forgot the mutual energy term \(\pm M i_1 i_2\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_W-3-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_mcq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans4</sans>
        <tans>ta_mcq</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mcq-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! All energy-increasing changes identified.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mcq-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all correct changes identified. Consider how each parameter affects the energy formula.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Q3 V4: Rectangular Core, Opposing Dots, Tight Coupling -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Midterm 2 - Q3V4: Coupled Circuits + Transformer (Rectangular Core, Opposing Dots, Tight Coupling)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Midterm 2 &ndash; Q3V4: Coupled Circuits + Transformer Design (Rectangular Core, Opposing Dots, Tight Coupling)</h3>
<p><em>Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><strong>Total: 10 points | Penalty per wrong attempt: 10%</strong></p>
<p><em>Inspired by Nilsson &amp; Riedel Ch.9 Problems 9.6&ndash;9.24.</em></p>
<hr />

<h4>Scenario</h4>
<p>Two magnetically coupled coils have self-inductances \(L_1 = 60\;\text{mH}\) and
\(L_2 = 15\;\text{mH}\), with coupling coefficient \(k = {@k_val@}\).</p>

<p>The coils are wound in <strong>opposite directions</strong> on a rectangular core. Current \(i_1\) enters the dotted terminal of Coil 1. Current \(i_2\) enters the <strong>undotted</strong> terminal of Coil 2.</p>

<p>[DIAGRAM: Physical winding diagram for Rectangular Core, Opposing Dots, Tight Coupling &mdash; to be added]</p>

<p>The coils carry DC currents \(i_1 = 2\;\text{A}\) and \(i_2 = 5\;\text{A}\).</p>

<hr />

<h4>Part (a) &mdash; Dot Configuration (1 point)</h4>
<p>Based on the winding description above, select the correct dot placement:</p>
<p>[[input:ans1]] &nbsp; [[validation:ans1]]</p>
<p><em>(This part will be upgraded to drag-and-drop when diagrams are available.)</em></p>

<h4>Part (b) &mdash; Mutual Inductance (2 points)</h4>
<p>Calculate the mutual inductance \(M\).</p>
<p>\(M = \) [[input:ans2]] \(\text{H}\) &nbsp; [[validation:ans2]]</p>
<p><em>Enter a number, e.g. <code>0.0456</code></em></p>

<h4>Part (c) &mdash; Stored Energy (2 points)</h4>
<p>Calculate the total magnetic energy stored in the coupled system.</p>
<p>Recall: \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\)</p>
<p>The sign of the mutual term depends on the dot convention and current directions.</p>
<p>\(W = \) [[input:ans3]] \(\text{J}\) &nbsp; [[validation:ans3]]</p>
<p><em>Enter a number, e.g. <code>0.567</code></em></p>

<h4>Part (d) &mdash; Energy Optimization (1 point)</h4>
<p><strong>Select ALL changes that would increase the total stored energy:</strong></p>
<p>[[input:ans4]] &nbsp; [[validation:ans4]]</p>

<h4>Part (e) &mdash; Ideal Transformer Comparison (4 points, manually graded)</h4>
<p>An engineer proposes modeling these coupled coils as an <strong>ideal transformer</strong> with
turns ratio \(a = \sqrt{L_1/L_2} = {@float(a_val)@}\).</p>
<ol>
  <li>Under what conditions is the ideal transformer approximation valid?</li>
  <li>What physical effects are lost in the ideal transformer model?</li>
  <li>For the given currents, what is the relationship between \(i_1\) and \(i_2\)
      under the ideal transformer assumption? Does this match the given values?</li>
</ol>
<p>[[input:ans5]] &nbsp; [[validation:ans5]]</p>
]]></text>
  </questiontext>
  <generalfeedback format="html">
    <text><![CDATA[
<h4>Solution for V4 (Rectangular Core, Opposing Dots, Tight Coupling)</h4>
<p><strong>Part (a):</strong> The dot configuration is opposing.</p>
<p><strong>Part (b):</strong> \(M = k\sqrt{L_1 L_2} = {@M_num@}\;\text{H}\)</p>
<p><strong>Part (c):</strong></p>
<p>\(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 - M i_1 i_2 = {@W_num@}\;\text{J}\)</p>
<p>The sign is <strong>-</strong> because the dot configuration is <strong>opposing</strong>.</p>
<p><strong>Part (d):</strong> The correct selections are {@correct_mcq@}.</p>
<p><strong>Part (e):</strong></p>
<ul>
  <li>The ideal transformer requires \(k \to 1\) (perfect coupling) and infinite inductances with finite ratio</li>
  <li>Lost effects: leakage flux, core losses, winding resistance, finite inductance (DC magnetization)</li>
  <li>Ideal transformer: \(i_1/i_2 = 1/a = {@float(1/a_val)@}\). The given currents \(i_1 = 2\) A and \(i_2 = 5\) A may not satisfy this ratio, showing the coils are not ideal.</li>
</ul>
]]></text>
  </generalfeedback>
  <defaultgrade>10</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2024032401</text></stackversion>
  <questionvariables>
    <text><![CDATA[
/* V4: Rectangular Core, Opposing Dots, Tight Coupling */
L1_val: 0.06;
L2_val: 0.015;
k_val: 0.45;
i1_val: 2;
i2_val: 5;

/* Mutual inductance */
M_val: k_val*sqrt(L1_val*L2_val);

/* Stored energy: W = 1/2*L1*i1^2 + 1/2*L2*i2^2 - M*i1*i2 */
W_val: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 - M_val*i1_val*i2_val;

/* Error values */
W_wrong_sign: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2 + M_val*i1_val*i2_val;
W_forgot_M: 1/2*L1_val*i1_val^2 + 1/2*L2_val*i2_val^2;

/* T-equivalent elements */
La_val: L1_val + M_val;
Lb_val: L2_val + M_val;

/* Turns ratio estimate */
a_val: sqrt(L1_val/L2_val);

/* Numerical versions */
M_num: float(M_val);
W_num: float(W_val);

/* Part (a) dot config MCQ */
dot_opts: random_permutation([[1, false, "Dots on the same side of both coils"], [2, true, "Dot on top of Coil 1, dot on bottom of Coil 2 (opposite winding)"], [3, false, "No dots are needed for this configuration"], [4, false, "Dots on opposite sides, but reversed from option B"]]);

/* Part (d) MCQ-MA */
mcq_opts: random_permutation([[1, false, "Decrease k"], [2, true, "Increase i_2"], [3, true, "Increase L_2"], [4, false, "Reverse i_1 direction (makes mutual term positive)"], [5, false, "Replace opposing dots with aiding dots"]]);
correct_mcq: {2,3};

/* Teacher answers */
ta_dot: dot_opts;
ta_M: float(M_val);
ta_W: float(W_val);
ta_mcq: correct_mcq;
]]></text>
  </questionvariables>
  <specificfeedback format="html">
    <text><![CDATA[
[[feedback:prt_dot]]
[[feedback:prt_M]]
[[feedback:prt_W]]
[[feedback:prt_mcq]]
]]></text>
  </specificfeedback>
  <questionnote><text>{@M_num@}, {@W_num@}, {@correct_mcq@}</text></questionnote>
  <questiondescription format="html"><text></text></questiondescription>
  <prtcorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text></prtcorrect>
  <prtpartiallycorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text></prtpartiallycorrect>
  <prtincorrect format="html"><text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text></prtincorrect>
  <multiplicativetries>0</multiplicativetries>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_dot</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_M</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.0456</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_W</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>0.567</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>checkbox</type>
      <tans>ta_mcq</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>essay</type>
      <tans></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_dot</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_dot-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct dot placement.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_dot-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. The dot marks the terminal where current entering produces flux in the reference direction. Check the winding sense.</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_M</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta_M</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_M-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct mutual inductance.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_M-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(M = k\sqrt{L_1 L_2}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_W</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>ta_W</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct stored energy.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_W-1-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_wrong_sign)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-2-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Wrong sign on the mutual term! Check whether the configuration is aiding or opposing. Aiding: +M, Opposing: -M.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_W-2-F</falseanswernote>
        <falsefeedback format="html"><text></text></falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>NumRelative</answertest>
        <sans>ans3</sans>
        <tans>float(W_forgot_M)</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_W-3-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>You computed only the self-energy terms and forgot the mutual energy term \(\pm M i_1 i_2\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_W-3-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Incorrect. \(W = \frac{1}{2}L_1 i_1^2 + \frac{1}{2}L_2 i_2^2 \pm M i_1 i_2\).</p>]]></text></falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_mcq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Sets</answertest>
        <sans>ans4</sans>
        <tans>ta_mcq</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mcq-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Correct! All energy-increasing changes identified.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mcq-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Not all correct changes identified. Consider how each parameter affects the energy formula.</p>]]></text></falsefeedback>
      </node>
    </prt>
</question>

<!-- ============================================================ -->
<!-- Companion Essay -->
<!-- ============================================================ -->
<question type="essay">
  <name><text>Midterm 2 - Q3: Upload Handwritten Work (Coupled Circuits)</text></name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Upload Your Handwritten Work for Q3</h3>
<p>Upload handwritten work for Question 3 (Coupled Circuits + Transformer Design).</p>
<h4>Include:</h4>
<ul>
  <li><strong>Part (a):</strong> Winding analysis and dot reasoning</li>
  <li><strong>Part (b):</strong> M calculation</li>
  <li><strong>Part (c):</strong> Energy calculation with sign reasoning</li>
  <li><strong>Part (e):</strong> Ideal transformer comparison and analysis</li>
</ul>
<h4>Instructions:</h4>
<ol>
  <li>Name and student ID on every page</li>
  <li>Dark pen on white paper</li>
  <li>Up to 2 files: PDF, JPG, JPEG, PNG</li>
</ol>
<p><em>Worth 0 points. Supports grading verification.</em></p>
<p><strong>Teacher notice:</strong> Set &ldquo;Allow attachments&rdquo; to 2+ in quiz settings.</p>
]]></text>
  </questiontext>
  <generalfeedback format="html"><text><![CDATA[<p>Thank you for uploading.</p>]]></text></generalfeedback>
  <defaultgrade>0</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <responseformat>noinline</responseformat>
  <responserequired>0</responserequired>
  <responsefieldlines>5</responsefieldlines>
  <attachments>2</attachments>
  <attachmentsrequired>1</attachmentsrequired>
  <filetypeslist>.pdf,.jpg,.jpeg,.png</filetypeslist>
  <graderinfo format="html">
    <text><![CDATA[
<h4>Instructor Notes</h4>
<ul>
  <li>Check dot reasoning from winding description</li>
  <li>Verify correct sign in energy formula (most common error)</li>
  <li>For Part (e): look for k&rarr;1 condition, mention of leakage flux, finite inductance</li>
</ul>
]]></text>
  </graderinfo>
  <responsetemplate format="html"><text></text></responsetemplate>
</question>

</quiz>